```python
import numpy as np
import pandas as pd
from numpy.testing import assert_equal

import warnings
warnings.filterwarnings("ignore")
```

You will be working with the Sushi Preference dataset. The complete documentation is at `/mnt/data/public/sushipref/sushi3-2016/README-en.txt`.

The relevant parts of the documentation are reproduced below. 

    ### sushi3b.5000.10.score #########################################

    - matrix style data separated by <sp>
    - each row corresponds to the user in the corresponding line of the file sushi3.udata
    - each column corresponds to the SUSHI in the item set B
    - using five-point-scale, 0:the most disliked, 4:the most preferred, -1:not rated
    
    ### sushi3b.udata #################################################

    matrix style format separated by TAB

    * each row represents one user that corresponds to the line in the .order or .score file
    * each column represents the following feature
    1.  the user ID
    2.  gender   0:male 1:female
    3.  age      0:15-19  1:20-29   2:30-39   3:40-49   4:50-59    5:60-
    4.  the total time need to fill questionnaire form
    5.  prefecture ID at which you have been the most longly lived until 15 years old
    6.  region ID at which you have been the most longly lived until 15 years old
    7.  east/west ID at which you have been the most longly lived until 15 years old
    8.  prefecture ID at which you currently live
    9.  regional ID at which you currently live
    10. east/west ID at which you currently live
    11. 0 if features 5 and 8 are equal; 1 otherwise

    ### sushi3.idata #################################################
    
    see also http://en.wikipedia.org/wiki/Sushi
    
    matrix style format separated by TAB
    
    * the row number corresponds to IDs of item set B
    * each column represents the following feature
    1. item ID
    2. name (in Japanese with Roman alphabets)
    3. style         0:maki      1:otherwise (see Wikipedia)
    4. major group   0:seafood   1:otherwise
       0 corresponds to the minor group nos 0--8.
    5. minor group
        0:aomono (blue-skinned fish)
        1:akami (red meat fish)
        2:shiromi (white-meat fish)
        3:tare (something like baste; for eel or sea eel)
        4:clam or shell
        5:squid or octopus
        6:shrimp or crab
        7:roe
        8:other seafood
        9:egg
       10:meat other than fish
       11:vegetables
    6. the heaviness/oiliness in taste, range[0-4] 0:heavy/oily
    7. how frequently the user eats the SUSHI, range[0-3] 3:frequently eat
    8. normalized price
    9. how frequently the SUSHI is sold in sushi shop, range[0-1] 1:the most frequently
    
    - features 3-5 are categorical, and features 6-9 are numerical

**Each cell must finish running within 5 minutes. To minimize memory usage, put all of your code inside the functions you are asked to create.**

## Problem 1

Create a function `get_utility` that will read the Sushi Preference dataset and returns the utility matrix as a pandas Data Frame where the columns are the sushi names and the labels are the user IDs. Represent unrated items as `np.nan`. Sort by user ID.

<!-- grade_id: cell-4b74833b7d52b781, points: 0, grade: false, solution: true -->
```python
def get_utility():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-901ffdd4e0df2136, points: 1, grade: true, solution: false -->
```python
df_utility = get_utility()
assert_equal(df_utility.shape, (5000, 100))
assert_equal(
    df_utility.columns[:10].tolist(),
    [
        "ebi",
        "anago",
        "maguro",
        "ika",
        "uni",
        "tako",
        "ikura",
        "tamago",
        "toro",
        "amaebi",
    ],
)
assert_equal(
    df_utility.index[:10].tolist(),
    [0, 1, 2, 3, 1000, 1001, 1003, 1004, 1005, 1006]
)
assert_equal(
    df_utility.iloc[:10, :10].to_numpy().tolist(),
    [
        [np.nan, np.nan, np.nan, np.nan, 3.0,
         np.nan, np.nan, 4.0, np.nan, np.nan],
        [np.nan, np.nan, np.nan, 4.0, np.nan,
         np.nan, 3.0, 2.0, 4.0, np.nan],
        [np.nan, np.nan, np.nan, np.nan, np.nan,
         np.nan, np.nan, 4.0, np.nan, np.nan],
        [np.nan, np.nan, 4.0, np.nan, np.nan,
         3.0, np.nan, 3.0, np.nan, 2.0],
        [np.nan, np.nan, np.nan, 2.0, np.nan,
         np.nan, np.nan, np.nan, np.nan, 3.0],
        [np.nan, np.nan, 4.0, np.nan, np.nan,
         4.0, 4.0, np.nan, np.nan, np.nan],
        [np.nan, np.nan, np.nan, np.nan, np.nan,
         2.0, np.nan, np.nan, np.nan, 3.0],
        [np.nan, 4.0, np.nan, np.nan, np.nan,
         3.0, np.nan, np.nan, np.nan, 4.0],
        [np.nan, np.nan, np.nan, np.nan, np.nan,
         np.nan, np.nan, np.nan, np.nan, 4.0],
        [2.0, np.nan, np.nan, np.nan, np.nan,
         3.0, np.nan, np.nan, np.nan, np.nan],
    ],
)
assert_equal(
    df_utility.iloc[:10, -10:].to_numpy().tolist(),
    [
        [
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
        ],
        [
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
        ],
        [np.nan, np.nan, np.nan, np.nan, np.nan,
         np.nan, np.nan, np.nan, np.nan, 3.0],
        [
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
        ],
        [
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
        ],
        [
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
        ],
        [
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
        ],
        [
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
            np.nan,
        ],
        [np.nan, np.nan, np.nan, np.nan, np.nan,
         np.nan, np.nan, np.nan, 1.0, np.nan],
        [np.nan, 4.0, np.nan, np.nan, np.nan,
         np.nan, np.nan, np.nan, np.nan, np.nan],
    ],
)
```

## Problem 2

Create a function `largest_liked_together_baskets` that returns the list of baskets of sushi names that are frequently liked together, where no basket in the list can be extended with another sushi and still be frequently liked together. A user likes a sushi if they rated it at least 3. A basket is frequently liked together if at least 0.2% of all users like every sushi in the basket. Consider only baskets with at least two sushi names. Return each basket as a list of sushi names sorted alphabetically, and sort the baskets by decreasing number of users who like the whole basket, breaking ties by comparing the baskets element by element.

<!-- grade_id: cell-99570d67392c29ba, points: 0, grade: false, solution: true -->
```python
def largest_liked_together_baskets():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e776cd5bad542eed, points: 1, grade: true, solution: false -->
```python
lltb = largest_liked_together_baskets()
assert_equal(len(lltb), 4341)
assert_equal(
    lltb[:10],
    [
        ["ebi", "ikura", "tako"],
        ["ika", "maguro", "tai"],
        ["amaebi", "maguro", "toro"],
        ["anago", "hamachi", "maguro"],
        ["hamachi", "maguro", "toro"],
        ["maguro", "tamago", "toro"],
        ["ebi", "hotategai", "tako"],
        ["ebi", "maguro", "tako"],
        ["ebi", "sarada"],
        ["hotategai", "sayori"],
    ],
)
```

## Problem 3

Create a function `assoc_rules` that returns the association rules using the same liked-sushi data as before. Each rule has one or more sushi in the antecedent and exactly one sushi in the consequent. Each rule must have a relative support of at least 0.2% and a confidence of at least 0.6. List each rule as a tuple of antecedent and consequent. Items in the antecedent should be sorted alphabetically. Sort the rules by decreasing lift, then by consequent alphabetically, then by comparing the antecedents element by element.

<!-- grade_id: cell-04a307e756fb75ba, points: 0, grade: false, solution: true -->
```python
def assoc_rules():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-aa89a6a97e1d3386, points: 1, grade: true, solution: false -->
```python
rules = assoc_rules()
assert_equal(len(rules), 40)
assert_equal(
    rules[:10],
    [
        (("maguro", "tako", "tobiko"), "shimaaji"),
        (("ebi", "ika", "mirugai"), "torigai"),
        (("basashi", "kanpachi"), "saba"),
        (("shako", "tako", "tamago"), "aji"),
        (("botanebi", "ebi", "engawa"), "samon"),
        (("sanma", "unagi"), "tamago"),
        (("kani", "shako", "tako"), "chu_toro"),
        (("kohada", "maguro", "tamago"), "tai"),
        (("amaebi", "tsubugai"), "hotategai"),
        (("aji", "shako", "tamago"), "tako"),
    ],
)
```

## Problem 4 [2 pts]

Provide three recommendations to the owner of a sushi restaurant based on the frequent itemset mining and association rule mining performed on the sushi dataset. Each recommendation should be supported by relevant FIM or ARM results.

## Problem 5

Create a function `user_recommend` that accepts a user ID and utility matrix then returns the 10 best sushi recommendations for that `user` by a user-based collaborative filtering recommender system using (up to) $k=50$ nearest neighbors. Consider only neighbors with at least 3 commonly rated items as the `user`. Follow the preprocessing and conventions used in class. Sort by decreasing estimated rating then by ascending order of the sushi name.

<!-- grade_id: cell-313de3d03533e2ec, points: 0, grade: false, solution: true -->
```python
def user_recommend(user, df_utility):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-3db555b4fc37fab4, points: 1, grade: true, solution: false -->
```python
assert_equal(
    user_recommend(1, df_utility),
    [
        "toro_samon",
        "gyusashi",
        "kappa_maki",
        "hotategai",
        "kanpachi",
        "chu_toro",
        "shirauo",
        "negi_toro_maki",
        "hamachi",
        "amaebi",
    ],
)
```

## Problem 6

Create a function `item_recommend` that accepts a user ID and utility matrix then returns the 10 best sushi ID recommendations for that `user` by an item-based collaborative filtering recommender system using (up to) $k=10$ nearest neighbors. Consider only neighbors with at least 3 commonly rated items as the `item`. Follow the preprocessing and conventions used in class. Sort by decreasing estimated rating then by ascending order of the sushi name.

<!-- grade_id: cell-188a75efd7e95aba, points: 0, grade: false, solution: true -->
```python
def item_recommend(user, df_utility):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d26cbf306b678975, points: 1, grade: true, solution: false -->
```python
assert_equal(
    item_recommend(1, df_utility),
    [
        "engawa",
        "hamachi",
        "karei",
        "negi_toro_maki",
        "suzuki",
        "tai",
        "chu_toro",
        "maguro",
        "negi_toro",
        "hiramasa",
    ],
)
```

## Problem 7

Create a function `cgd_recommend` that accepts a user ID and utility matrix then returns the 10 best sushi recommendations for that `user` by a CGD-based recommender system using 20 latent components. Stop when the RMSE is at most 1.03. Sort the recommendations by decreasing estimated rating then by ascending order of the sushi name.

<!-- grade_id: cell-276caf4456eba059, points: 0, grade: false, solution: true -->
```python
def cgd_recommend(user, df_utility):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-126d31372f0797d9, points: 1, grade: true, solution: false -->
```python
assert_equal(
    cgd_recommend(0, df_utility),
    [
        "chu_toro",
        "toro",
        "negi_toro",
        "maguro",
        "negi_toro_maki",
        "tarabagani",
        "amaebi",
        "kurumaebi",
        "ebi",
        "tekka_maki",
    ],
)
```

## Problem 8 [2 pts]

Compare the recommendations from problems 5 to 7. Which do you think is the best set of recommendations? Justify your answer.