<img src="images/banner.png" style="width: 100%;">

# Introduction to `nbgrader`

For this course, you will be submitting notebooks via Jojie (Jupyterhub). This platform, which runs on [nbgrader](https://github.com/jupyter/nbgrader), provides a more seamless fetching and submitting of assignments for students. For teachers, it provides them with more convenient grading interfaces which allow both manual and automated grading. 

It is important that students properly submit notebooks via this platform to avoid risking lower scores due to unintended confusion on the part of the teacher and automated grader. In this notebook, you will answer and submit notebooks via this platform.

## Problem 1 [1 pt]

This is a manually graded problem that requires a valid Python code for the solution. The answer cell is the succeeding cell and contains the following code snippet:

    # YOUR CODE HERE
    raise NotImplementedError()
    
To answer, comment out or delete the `raise NotImplementedError()` statement. You may also remove the entire content of the cell and replace it with your answer. **Do not** copy, delete or move that cell around.

<div class='alert alert-info'>
    
**Manually Graded Problem (Code)**

Complete the function below

</div>

<!-- grade_id: cell-121de2cf60e780e2, points: 1, grade: true, solution: true -->
```python
def get_fav_color():
    """Return my favorite color"""
    fav_color = 'Blue'
    return fav_color
```

## Problem 2 [1 pt]

The required answer may also be in the form of text. In this case, the answer cell would be a markdown cell that contains the placeholder text: `YOUR ANSWER HERE`. Just like in other types of answer cells, you should restrict your answer within the cell and avoid adding or moving around cells.

<div class='alert alert-info'>
    
**Manually Graded Problem (Text)**

Describe yourself in one sentence

</div>

<!-- grade_id: cell-c3a8ac1bda65576a, points: 1, grade: true, solution: true -->
I like gambling.

## Problem 3 [3 pts]

Sometimes the problem requires you to perform a task. In that case, there is no answer cell following the statement of the problem. You are instead expected to add as many cells as needed after the cell to accomplish the required task. Task problems are manually graded.

<div class='alert alert-info'>

**Manually Graded Task**

Answer the following in separate markdown cells:
* What's your motto in life?
* If you were a fictional character, who would you be and why?
* If you were trapped in an island, why would you spend your time coding in Python?

</div>

Gambling is life!

If I was a fictional character, I would be the trailblazer from Honkai: Star Rail.

If I am trapped on an island, I would spend my time coding in Python because I would be programming a flying machine to get out.

## Problem 4 [1 pt]

This is an automatically graded problem so it is very important that you properly submit your answer. 

<div class='alert alert-info'>
    
**Autograded Problem**

What is the area of a square with sides of 2 units?

</div>

<!-- grade_id: cell-9b9c7c8dd8ea5fb3, points: 0, grade: false, solution: true -->
```python
def area_2x2():
    area = 0
    side = 2
    area = side*side
    return area
```

Autograded answers are usually followed by one or more cells containing automated tests. You **cannot** delete or edit them. **Do not** copy and paste these cells.

There can be hidden tests so the tests that you see may not be complete.

<!-- grade_id: cell-d67d713acad56093, points: 1, grade: true, solution: false -->
```python
from numbers import Number
from numpy.testing import assert_almost_equal

area = area_2x2()
assert isinstance(area, Number)
```

## Problem 5 [1 pt]

<div class='alert alert-info'>
    
**Autograded Problem (multiple tests)**

Complete the function below

</div>

<!-- grade_id: cell-9626ae57a69e5d87, points: 0, grade: false, solution: true -->
```python
def always_true(a):
    """Return True always"""
    return True
```

Autograded answers may have more than one tests cell.

<!-- grade_id: cell-459839ccb2fd81cb, points: 0.5, grade: true, solution: false -->
```python
assert always_true(1) is True
assert always_true(True) is True
```

The tests in a tests cell may all be hidden as well. In that case, leave it alone as usual.

<!-- grade_id: cell-994da6990cc81fce, points: 0.5, grade: true, solution: false -->
```python
ans = always_true(False)
```

This ends the introductory exercise. You may now click `Validate` to check whether the notebook passes all visible tests. Validation checks notebook metadata and runs the notebook from top to bottom. However, it has a 30-second timeout for each cell: a cell running for more than that time will be considered as an error. If you are expecting that a cell runs by more than that amount, you may ignore the error or just run the notebook from top to bottom instead. 

Save this notebook then submit it. **Do not** rename this notebook. Make sure you don't submit past the deadline otherwise you'll automatically get deductions.

<img src="images/banner-down.png" style="width: 100%;">